# Exercise 12 – Feature engineering 

**General Instructions:**

- Review the weekly course material (lectures, readings, slides, etc.) before starting the exercises.
- Complete all assigned exercises independently before the Q&A session.
- Please use Julia Version 1.12.x to ensure compatibility.
- Please only write between the `#--- YOUR CODE STARTS HERE ---#` and `#--- YOUR CODE ENDS HERE ---#` comments.

```julia
using LinearAlgebra
using Plots
using HDF5
using MAT
using Flux
using Serialization
using Random
using Statistics
using MLUtils 
```

## Task 1 – Extended Dynamic Mode Decomposition using a dictionary
Learn a linear dynamical system from trajectory data via the Dynamic Mode Decomposition. As the example, we will use the Duffing oscillator with a two-dimensional state $x\in\mathbb{R}^2$. The function for the right-hand side is given below

```julia
function rhs(x)
    ɑ, β, δ = -1, 1, 0.1
    xdot = [x[2], - δ * x[2] - ɑ * x[1] - β * x[1]^3]
    return xdot
end
```

a) Create training data $X$ and $X’$. To this end, draw $N=1000$ initial conditions randomly and uniformly from the rectangle $[-2,2] \times [-2,2]$ and perform one time step using the explicit Euler scheme with a time step of $h = 0.1$. 

<!-- grade_id: cell-90d5f1b7fcfc1950, points: 0, grade: false, solution: true -->
```julia
X = nothing
X_prime = nothing

### BEGIN SOLUTION

### END SOLUTION
```

<!-- grade_id: cell-9b4839e10ebab21d, points: 1, grade: true, solution: false -->
```julia
@assert size(X) == (2, 1000)
@assert size(X_prime) == (2, 1000)

### BEGIN TESTS

#Check if sample points are in range
for value in X
    @assert value >= -2 && value <= 2
end

#Check if X_prime is calculated correctly
for i in 1:1000
    @assert isapprox(X_prime[:, i], X[:, i] + 0.1 * rhs(X[:, i]))
end
### END TESTS
```

b) Use DMD in its standard form to learn a linear system with $A\in\mathbb{R}^{2\times 2}$. Simulate the dynamics over $m=100$ time steps using the initial condition $x_0=[1.0, -1.0]$. Report on the RMSE between the created trajectory and a true trajectory of the system (same initial condition, explicit Euler integration).

Note: To avoid confusion regarding the RMSE on multiple dimensions, take this one:
$ \text{RMSE} = \sqrt{ \frac{1}{T} \sum_{t = 1}^{T} \sum_{i = 1}^{d} (X_{t, i} - Y_{t, i})^2 } $

<!-- grade_id: cell-38b95b928edfdbfa, points: 0, grade: false, solution: true -->
```julia
x0 = [1.0, -1.0]

A = nothing
dmd_trajectory = nothing
euler_trajectory = nothing
RMSE = nothing

### BEGIN SOLUTION

### END SOLUTION

```

c) Implement a dictionary of radial basis functions (RBFs), i.e.,  
$$ 
\varphi_c(x) = \exp(- \gamma \| x - c \|_2^2 ),
$$ 
$$ 
z = \Psi(x) = [\varphi_{c_1}(x), \ldots, \varphi_{c_r}(x)],
$$ 
where the $c_i\in\mathbb{R}^2$ are the centers of the individual RBFs and $\gamma>0$ is a hyperparameter determining the width of the RBF 

<!-- grade_id: cell-7f0011b600a0ae1e, points: 0, grade: false, solution: true -->
```julia
# Function to compute the lifted state using a dictionary of RBFs
function rbf_dictionary(x, centers, gamma=1.5)
    ### BEGIN SOLUTION
    
    ### END SOLUTION
end
```

<!-- grade_id: cell-c92001473d595796, points: 1, grade: true, solution: false -->
```julia
@assert isa(rbf_dictionary, Function)
@assert length(rbf_dictionary([1, 1], [[2, 2], [3, 3]])) == 2
@assert length(rbf_dictionary([1, 1], [[2, 2], [3, 3], [4, 4]])) == 3

### BEGIN TESTS
@assert isapprox(rbf_dictionary([1, 1], [[2, 2], [3, 3]]), [0.049787068367863896, 6.144212353328188e-6])
@assert isapprox(rbf_dictionary([0, 0], [[2, 2], [2, -2], [-2, 2], [-2, -2]]), [6.144212353328188e-6, 6.144212353328188e-6, 6.144212353328188e-6, 6.144212353328188e-6])
### END TESTS
```

d) Use the extended version of DMD (aka EDMD), where you find a linear system for the lifted state $\Psi(x)=z\in\mathbb{R}^r$ instead of $x\in\mathbb{R}^2$ (use the trajectory from a) ). To this end, introduce $r=100$ centers on an equidistant $10 \times 10$ grid in the area $[-2, 2] \times [-2, 2]$, with $\gamma = 1$. Train the corresponding matrix $A\in\mathbb{R}^{100 \times 100}$. 

<!-- grade_id: cell-541dfb172c231815, points: 0, grade: false, solution: true -->
```julia
centers = nothing
Z = nothing
Z_prime = nothing
A_edmd = nothing

### BEGIN SOLUTION

### END SOLUTION
```

e) Implement a projection operation that maps from $z$ to $x$, such that you can reconstruct the original state $x$ from predictions of the lifted state $z$. This step can be realized by a linear mapping $x = P z$, where the projection matrix $P\in\mathbb{R}^{2\times r}$ can be determined using linear regression. 

<!-- grade_id: cell-969fbc2f15683cf4, points: 0, grade: false, solution: true -->
```julia
P = nothing

### BEGIN SOLUTION

### END SOLUTION
```

f) Repeat the experiment from b), but now compare the original trajectory to the one obtained using EDMD and the consecutive projection step to recover $x$. 

```Hint: Don't forget to transform the initial condition accordingly```

<!-- grade_id: cell-9e2cda54252dbd95, points: 0, grade: false, solution: true -->
```julia
edmd_trajectory_lifted = nothing
edmd_trajectory_projected = nothing
error_edmd = nothing

### BEGIN SOLUTION

### END SOLUTION
```

## Task 2 – Comparing SVD and Autoencoders
In this exercise, we want to compare the capabilities of the Singular Value Decomposition (SVD) and an autoencoder in terms of data compression of high-dimensional trajectory data. The following data set $X\in\mathbb{R}^{101 \times 6001}$ has been created by simulating the Burgers PDE (https://en.wikipedia.org/wiki/Burgers%27_equation) on an $n = 101$-dimensional spatial grid for $N=6000$ time steps. 

a) Use the SVD to compress the $n$-dimensional state into a much lower dimension $k\ll n$. Select the smallest value for $k$ for which the reconstruction error between the original data $x$ and the reconstructed data $\tilde{x}$ is less than 0.1%, i.e. 
$$ 
\frac{ \sum_{i=1}^N \|x_i – \tilde{x}_i \|_2^2 }{ \sum_{i=1}^N \|x_i \|_2^2 } < 0.001. 
$$ 

```julia
#Data loading
file = matopen("burgers.mat")
X = read(file, "u")
close(file)
```

<!-- grade_id: cell-b2dc2b48b80e38a6, points: 0, grade: false, solution: true -->
```julia
function loss_function(original, reconstructed)
    ### BEGIN SOLUTION
   
    ### END SOLUTION
end
```

<!-- grade_id: cell-6dce53992aba9a2a, points: 0.5, grade: true, solution: false -->
```julia
@assert isa(loss_function, Function)

### BEGIN TESTS
@assert isapprox(0.0035087719298245615, loss_function([1 2 3; 4 5 6; 7 8 9], [1 2 3; 4 5 6; 7 8 10]))
### END TESTS
```

<!-- grade_id: cell-59d3100f355f9745, points: 0, grade: false, solution: true -->
```julia
k = nothing
X_reconstructed_k = nothing

### BEGIN SOLUTION


### END SOLUTION
```

```julia
#Random SVD sample comparison
random_index = rand(1:size(X,2))
x = range(0.1,10.1, length = 101)
println("Original Image")
Plots.plot(x,X[:, random_index])
println("Reconstructed Image")
display(Plots.plot!(x,X_reconstructed_k[:, random_index]))
```

b) Now compare this to an autoencoder architecture of your choice. The only constraints that should be respected are: 
- Use only fully connected feed-forward layers (i.e., no convolutions) 
- the bottleneck layer connecting the encoder and the decoder (i.e., the smallest layer whose latent state is the compressed state) has to have dimension at most $r=5$. 
This task is fulfilled if your trained architecture can satisfy the error criterion from a) 

<!-- grade_id: cell-0cda242b711bf9c7, points: 0, grade: false, solution: true -->
```julia
#define your autoencoder
autoencoder = nothing

### BEGIN SOLUTION

### END SOLUTION

```

<!-- grade_id: cell-b42956b706417b65, points: 0, grade: true, solution: false -->
```julia
@assert isa(autoencoder, Flux.Chain)
```

<!-- grade_id: cell-1779d9df10c9d7de, points: 0, grade: false, solution: true -->
```julia
# Train your autoencoder.

### BEGIN SOLUTION

### END SOLUTION
```

```julia
#Random autoencoder sample comparison
x = range(0.1,10.1, length = 101)
random_index = rand(1:size(X,2))
original = X[:,random_index]
reconstruction = autoencoder(original)
Plots.plot(x,original)
display(Plots.plot!(x,reconstruction))
```