# Exercise 13: SINDy 

**General Instructions:**

- Review the weekly course material (lectures, readings, slides, etc.) before starting the exercises.
- Complete all assigned exercises independently before the Q&A session.
- Please use Julia Version 1.12.x to ensure compatibility.
- Please only write between the `#--- YOUR CODE STARTS HERE ---#` and `#--- YOUR CODE ENDS HERE ---#` comments.


In this exercise, we will identify a partial differential equation (PDE) from data. This means that we use the SINDy algorithm with a dictionary that contains typical terms as they can be found in PDEs. For a detailed introduction, see the following article: https://www.pnas.org/doi/10.1073/pnas.1517384113

Note: To simplify the notation, we are going to use the following abbreviations: 

$$
x_t = \frac{\partial x}{\partial t}, \quad x_{ss} = \frac{\partial^2 x}{\partial s^2}, \quad \ldots 
$$ 

The goal is to identify a PDE of the form 

$$
x_t = f(x, x_s, x_{ss}, x^2, x x_s, x_s x_{ss}, \ldots),
\qquad
\text{(PDE)}
$$ 

where the highest-order of derivatives in space is 2, i.e., $x_{ss}$.

**Use the given data matrix $\mathbf X\in\mathbb{R}^{N \times n}$ for the identification, where $N=26$ is the number of time steps, and $n=51$ is the number of points in space.**
We import the matrix as `X` below:

```julia
import Random
using MAT
using LinearAlgebra
file = matopen("pde.mat")
X = read(file, "pde") |> transpose  # we flatten the matrix later on, so orientation does not really matter
                                    # `transpose` makes it adhere to the article
close(file)

dt = 0.005
ds = 0.1
```

![](pde_visualization.png)

## Exercise a)

a) To set up the dictionary $\mathbf \Psi(x)$, we first need to calculate the partial derivatives of the state $x$ with respect to time and space, i.e., $x_t$, $x_s$ and $x_{ss}$.  
Compute these numerically, using central differences. Do not use an extra package for this, but rather do it “manually”, with discretization sizes `dt` and `ds`. 
Store the derivative values appropriate matrices $\mathbf X_t$, $\mathbf X_s$ and $\mathbf X_{ss}$.
Also note that you need only to compute the derivatives for interior points.
E.g., the matrix $\mathbf X_t$ will have size ${(N-2) \times (n-2)}$. 

<!-- grade_id: cell-90d8768c915d6fb3, points: 0, grade: false, solution: true -->
```julia
Xt = nothing
Xs = nothing
Xss = nothing

### BEGIN SOLUTION
N, n = size(X) # n = points in space, N = points in time

## preallocate result arrays:
Xt = zeros(N - 2, n - 2)
Xs = similar(Xt)
Xss = similar(Xt)

for i in 2:(N-1)
    _i = i - 1
    for j in 2:(n-1)
        _j = j - 1
        ## space data is sorted along columns:
        Xs[_i, _j] = (X[i, j+1] - X[i, j-1]) / (2 * ds)
        Xss[_i, _j] = (X[i, j+1] - 2 * X[i, j] + X[i, j-1]) / (ds^2)
        ## time data is sorted along colums:
        Xt[_i, _j] = (X[i+1, j] - X[i-1, j]) / (2 * dt)
    end
end
### END SOLUTION
```

## Exercise b)

b) Flatten the data matrices `X`, `Xt`, etc. into long vectors of size $ (N-2) (n-2)$.
For `X`, remove boundary points beforehand. 

<!-- grade_id: cell-a514a9ac35feb0ae, points: 0, grade: false, solution: true -->
```julia
X_flat = nothing
Xt_flat = nothing
Xs_flat = nothing
Xss_flat = nothing

### BEGIN SOLUTION
X_reduced = X[2:end-1, 2:end-1]; # note: this copies. we could use `@view` or `view` instead.

## Flattening the matrices into long vectors
## A few different ways to do this:
X_flat = reshape(X_reduced, :); # a view into `X_reduced`
Xt_flat = reshape(Xt, (N - 2) * (n - 2)); # a view, but length given explicitely
Xs_flat = Xs[:]; # a copy
Xss_flat = vec(Xss); # a view
### END SOLUTION
```

## Exercise c)

c) Create a feature matrix $\mathbf \Psi$ of possible right-hand side terms using all terms up to order two.
That is, the column-layout is as follows:
$$ \mathbf \Psi = [
    \mathbf 1, 
    \mathbf X, 
    \mathbf X_s, 
    \mathbf X_{ss}, 
    \mathbf X^2, 
    \mathbf X \cdot \mathbf X_s, 
    \mathbf X \cdot \mathbf X_{ss}, 
    \mathbf X_s^2, 
    \mathbf X_s \cdot \mathbf X_{ss}, 
    \mathbf X_{ss}^2]
    .
$$ 


<!-- grade_id: cell-140bd9d8a3ecfd82, points: 0, grade: false, solution: true -->
```julia
function Psi(X_flat, Xs_flat, Xss_flat)
    ### BEGIN SOLUTION
    N = size(X_flat, 1)

    psi = zeros(N, 10)
    psi[:, 1] .= 1

    psi[:, 2] .= X_flat
    psi[:, 3] .= Xs_flat
    psi[:, 4] .= Xss_flat

    psi[:, 5] .= X_flat .* X_flat
    psi[:, 6] .= X_flat .* Xs_flat
    psi[:, 7] .= X_flat .* Xss_flat
    psi[:, 8] .= Xs_flat .* Xs_flat
    psi[:, 9] .= Xs_flat .* Xss_flat
    psi[:, 10] .= Xss_flat .* Xss_flat
    ### END SOLUTION
    return psi
end;
```

## Exercise d)

d) Use the thresholding algorithm from the lecture to identify the PDE that created the data set $\mathbf X$.

First, assume that $\mathbf Y$ is a column vector with target data and $\mathbf Z$ is a basis matrix of suitable size.
Further, let $\mathbf w$ be the coefficient vector that we want to find, and $λ$ a real regularization parameter.
With LASSO we can solve
$$
\min_{\mathbf w \in ℝ^q} 
\left(
\underbrace{\| \mathbf Y - \mathbf Z \mathbf w \|_2^2}_{\mathcal L_{\text{reg}}(\mathbf w)} 
+
\underbrace{λ \|\mathbf w\|_1}_{\mathcal L_{\text{sparse}}(\mathbf w)}
\right)
.
$$
In this notebook, we takle LASSO by means of **coordinate descent**.
That is, we cyclically sweep over the parameters (i.e., the entries in $\mathbf w$) to optimize the combined loss function with respect to this single parameter whilst keeping all other weights fixed.
The optimal $w_j$ can be found by ensuring that $0$ is contained in the subgradient of the loss function.

It can be shown that
$$
\frac{\partial \mathcal L_{\text{reg}}}{\partial w_j} = 
2
\left(
-~
\underbrace{
\sum_{i=1}^{N_{\mathbf Y}}
\left(
    z_{i,j} 
        \left(
            y_i - \sum_{k\ne j} z_{i,k} w_k
        \right)
\right)
}_{ρ_j}
+
w_j
\underbrace{
\sum_{i=1}^{N_{\mathbf Y}}
   z_{i,j}^2
}_{φ_j}
\right)
=
2(-ρ_j + w_j φ_j)
,
$$
and
$$
\frac{\partial \mathcal L_{\text{sparse}}}{\partial w_j}
=
\begin{cases}
    -λ, &\text{if $w_j < 0$,}\\
    [-λ, λ] & \text{if $w_j = 0$},\\
    λ  &\text{if $w_j > 0$.}
\end{cases}
$$
From this, three cases for finding the optimal value $w_j$ are derived:
$$
w_j =
\begin{cases}
    \frac{ ρ_j + \frac{λ}{2} }{φ_j} & \text{if $ρ_j < - \frac{λ}{2}$,}\\
    0 & \text{ if $ρ_j \in \left[-\frac{λ}{2}, \frac{λ}{2}\right]$,}\\
    \frac{ ρ_j - \frac{λ}{2} }{φ_j} & \text{if $ρ_j > \frac{λ}{2}$.}
\end{cases}
$$

### Task 1

Complete the following functions to
* compute the values $ρ_j$ and $φ_j$ given vector $\mathbf Y$ and matrix $\mathbf Z$ and index $j$.
* return the optimal value $w_j$, if we addionally provide $λ>0$.

<!-- grade_id: cell-6de788356a4729b4, points: 0, grade: false, solution: true -->
```julia
function calc_rho_phi(Z::AbstractMatrix, Y::AbstractVector, w::AbstractVector, j::Integer)
    rho_j = phi_j = nothing

    ### BEGIN SOLUTION
    N = size(Z, 1)
    q = length(w)

    rho_j = 0
    for i in 1:N
        zw_ij = Z[i, :]'w - Z[i, j] * w[j]
        rho_j += Z[i, j] * (Y[i] - zw_ij)
    end

    phi_j = sum(Z[:, j] .^ 2)
    
    ### END SOLUTION
    return rho_j, phi_j
end;
```

<!-- grade_id: cell-543b0f1b1f0ee71a, points: 0.5, grade: true, solution: false -->
```julia
let
    Random.seed!(1234)
    
    Z = rand(3, 2)
    w = rand(2)
    Y = rand(3)
    
    rho, phi = calc_rho_phi(Z, Y, w, 1)
    @assert rho isa Real
    @assert phi isa Real
    
    @assert isapprox(rho, 0.24372423119974812; rtol=1e-4)
    @assert isapprox(phi, 0.4554981037149115; rtol=1e-4)
    
    ### BEGIN TESTS
    Random.seed!(1234)
    
    Z = rand(3, 2)
    w = rand(2)
    Y = rand(3)
    rho, phi = calc_rho_phi(Z, Y, w, 2)
    @assert isapprox(rho, 0.3894592987045164; rtol=1e-4)
    @assert isapprox(phi, 1.079800014004892; rtol=1e-4)
    ### END TESTS
end
```

<!-- grade_id: cell-c500fff9c3ef0d61, points: 0, grade: false, solution: true -->
```julia
function opt_wj(Z::AbstractMatrix, Y::AbstractVector, w::AbstractVector, j::Integer, λ::Real)
    w_j_opt = nothing
    
    ### BEGIN SOLUTION
    N = size(Z, 1)
    ρj, φj = calc_rho_phi(Z, Y, w, j)
    if ρj < -λ / 2
        w_j_opt = (ρj + λ / 2) / φj
    elseif ρj > λ / 2
        w_j_opt = (ρj - λ / 2) / φj
    else
        w_j_opt = 0
    end
    ### END SOLUTION
    
    return w_j_opt
end;
```

<!-- grade_id: cell-5a4d9c9079ad00cb, points: 0.5, grade: true, solution: false -->
```julia
let
    Random.seed!(1234)
    
    Z = rand(3, 2)
    w = rand(2)
    Y = rand(3)
    w_j_opt = opt_wj(Z, Y, w, 1, 0.1)
    @assert w_j_opt > 0
    
    ### BEGIN TESTS
    w_j_opt = opt_wj(Z, Y, w, 1, 0.5)
    @assert iszero(w_j_opt)
    
    Random.seed!(4321)
    
    Z = rand(3, 2)
    w = rand(2)
    Y = rand(3)
    w_j_opt = opt_wj(Z, Y, w, 1, 0.1)
    @assert w_j_opt < 0
    ### END TESTS
end
```

### Task 2 
Implement LASSO with Coordinate Descent by completing the cell below.

In iteration $k$, perform coordinate descent for all entries.
Stop, if $\|\mathbf w^{(k)} - \mathbf w^{(k-1)}\|_2 \leq ε_{\text{rel}} \|\mathbf w^{(k-1)}\|_2$ or if the maximum number of iterations is reached.

<!-- grade_id: cell-d6610a89304982b3, points: 0, grade: false, solution: true -->
```julia
function LASSO!(
        ## modify `w` in place
        w::AbstractVector,
        ## do not modify these:
        Z::AbstractMatrix, Y::AbstractVector, λ::Real;
        ## parameters:
        tol_rel = 1e-5,
        max_iter = 100_000
)
    ### BEGIN SOLUTION

    ## preallocate vector `w_prev` check convergence
    w_prev = similar(w)

    # count number of loops
    for k = 1 : max_iter
        ## copy old values
        w_prev .= w
        
        for j = eachindex(w)
            ## set optimal value for coordinate `j`
            w[j] = opt_wj(Z, Y, w, j, λ)
        end
        
        ## check stopping
        if norm(w_prev .- w, 2) <= tol_rel * norm(w_prev, 2)
            #@info "Stopping because relative stepsize is small."
            break
        end
    end
    ### END SOLUTION
    
    return w
end;
```

<!-- grade_id: cell-80b1204ed996ccc7, points: 1, grade: true, solution: false -->
```julia
let
    ## quickly generate a testset for polynomial fitting
    f(x) = sum( x.^2 )
    X = -2 .+ 4 .* rand(2, 10)
    q = 6
    Z = [
        ones(10);; X[1, :];; X[2, :];; X[1, :] .* X[2, :];; X[1,:] .^2;; X[2, :].^2 ]
    Y = mapreduce(f, vcat, eachcol(X))
    ## initialize parameter vector
    w = Vector{Float64}(undef, q)
    ## call LASSO!
    LASSO!(w, Z, Y, 1e-3; tol_rel = 1e-6)
    
    ## if the algorithm did nothing that is completely weird, the parameters should not have large magnitude:
    @assert all( abs(wj) <= 2 for wj = w)
    
    ### BEGIN TESTS
    ## see, if we actually identified the equations
    @assert all( abs(w[j]) <= 0.01 for j = 1:4 )
    @assert all( abs(1 - w[j]) <= 0.01 for j = 5:6 )
    ### END TESTS
end
```

### Task 3 
Finally, implement the alternating thresholding scheme.
To this end, apply `LASSO!` to those entries of $\mathbf w$ indexed by `I_nz`, the **n**on-**z**ero entries.
For initialization, use the ridge regression solution
$$
\mathbf w_0 = \left(\mathbf Z^\intercal \mathbf Z + λ \mathbf I\right)^{-1} \mathbf Z^\intercal \mathbf Y.
$$
In iteration $\ell$, first update `I_nz`. 
An entry of $\mathbf w^{(\ell)}$ is deemed to be zero, if its absolute value is below the threshold parameter `th`.
In addition to setting `I_nz`, also actually set the corresponding entries of $\mathbf w^{(\ell)}$ to exactly equal `0`.  
Then call `LASSO!` with stopping parameters `max_iter_lasso` and `rel_tol_lasso`.
Beware! We have implemented `LASSO!` to modify the parameter vector in-place.
Be sure to use views or perform copies accordingly!!

Again, implement a relative stopping criterion and respect the keyword-argument `max_iter`.

<!-- grade_id: cell-b7998be95cbf5b7b, points: 0, grade: false, solution: true -->
```julia
function alternating_thresholding_lasso(
        Z::AbstractMatrix, Y::AbstractVector, λ::Real;
        max_iter = 1000,
        tol_rel = 1e-4,
        th = 5e-3,
        max_iter_lasso = max_iter,
        tol_rel_lasso = 1e-3
)
    @info """
    Parameters are
      max_iter = $(max_iter),
      tol_rel  = $(tol_rel),
      th       = $(th),
      max_iter_lasso = $(max_iter_lasso),
      tol_rel_lasso = $(tol_rel_lasso)"""
    
    w = nothing
    ### BEGIN SOLUTION
    
    ## initialize `w` with ridge regression weights
    w = (Z' * Z + λ * I) \ Z' * Y
    
    ## pre-allocate array for previous values:
    w_prev = similar(w)
    
    ## mask for nonzero entries
    I_nz = ones(Bool, length(w))

    for l = 1:max_iter
        ## copy old values for stopping
        w_prev .= w
    
        ## update nonzero index
        for j = eachindex(I_nz)
            !I_nz[j] && continue
            if abs(w[j]) < th
                I_nz[j] = false
                w[j] = 0
            end
        end
        ## views for `LASSO!`
        w_nz = @view(w[I_nz])
        Z_nz = @view(Z[:, I_nz])
        
        ## call `LASSO!` to optimize `w_nz` (and thereby `w`):
        LASSO!(w_nz, Z_nz, Y, λ; tol_rel = tol_rel_lasso, max_iter = max_iter_lasso)
        
        ## check stopping
        if norm(w_prev .- w, 2) <= tol_rel * norm(w_prev, 2)
            @info "Stopping because outer relative stepsize is small."
            break
        end
    end

    ### END SOLUTION
    return w
end
```

<!-- grade_id: cell-6b01cdfcc9952d2e, points: 1, grade: true, solution: false -->
```julia
let
    ## quickly generate a testset for polynomial fitting
    f(x) = sum( x.^2 )
    X = -2 .+ 4 .* rand(2, 10)
    q = 6
    Z = [
        ones(10);; X[1, :];; X[2, :];; X[1, :] .* X[2, :];; X[1,:] .^2;; X[2, :].^2 ]
    Y = mapreduce(f, vcat, eachcol(X))
    ## call algorithm
    w = alternating_thresholding_lasso(
        Z, Y, 0.001;
        max_iter = 1000,
        tol_rel = 1e-3,
        th = 5e-3,
        max_iter_lasso = 1000, # note the reduced number of lasso iterations
        tol_rel_lasso = 1e-2   # and the coarser criterion
    )
    
    ## if the algorithm did nothing that is completely weird, the parameters should not have large magnitude:
    @assert all( abs(wj) <= 2 for wj = w)
    
    ### BEGIN TESTS
    ## see, if we actually identified the equations
    @assert all( abs(w[j]) <= 0.01 for j = 1:4 )
    @assert all( abs(1 - w[j]) <= 0.01 for j = 5:6 )
    ### END TESTS
end
```

### Task 4 
Determine, which vectors or matrices should be used to identify the equations governing $\text{(PDE)}$.
Do any reshaping that is still required.
Adhere to the column order used in previous exercises.
Call `alternating_thresholding_lasso` with $λ = 10^{-3}$ and parameters
```
max_iter = 1000,
tol_rel = 1e-3,
th = 5e-3,
max_iter_lasso = 1000,
tol_rel_lasso = 1e-3
```
Define `w` as the optimal sparse coefficient vector.

<!-- grade_id: cell-462549fc00fe355c, points: 1, grade: false, solution: true -->
```julia
lambda = w = Y = Z = nothing
### BEGIN SOLUTION
Y = vec(Xt_flat)
Z = Psi(X_flat, Xs_flat, Xss_flat)
lambda = 1e-3
w = alternating_thresholding_lasso(Z, Y, lambda; 
    max_iter = 1000,
    tol_rel = 1e-3,
    th = 5e-3,
    max_iter_lasso = 1000,
    tol_rel_lasso = 1e-3
);
### END SOLUTION
```

## Exercise e)

e) What is the name of the equation that you have found?
Assign one of the following values to `pde`:

1. The Kuramoto-Sivashinski equation 
2. The Keller-Segel model for chemotaxis 
3. The heat equation 
4. The Burgers equation 
5. The Navier-Stokes equations 
6. The wave equation 
7. The Poisson equation 
8. The Schrödinger equation 


<!-- grade_id: cell-fd9903bc71d6175e, points: 0, grade: false, solution: true -->
```julia
pde = nothing
### BEGIN SOLUTION
pde = 4;
### END SOLUTION
```