**General Instructions:**

- Review the weekly course material (lectures, readings, slides, etc.) before starting the exercises.
- Complete all assigned exercises independently before the Q&A session.
- Please use Julia Version 1.10.x to ensure compatibility.
- Please only write between the `#--- YOUR CODE STARTS HERE ---#` and `#--- YOUR CODE ENDS HERE ---#` comments.

# Exercise Sheet 09

## Recap

### First-Order Necessecary Conditions

The overall goal of this lecture series is to get a
feeling for solving the general nonlinear problem
$$
\begin{aligned}
&\min_{\mathbf w\in ℝ^q}
  L(\mathbf w)
      &\text{subject to}\\
&c_i(\mathbf w) = 0, &∀i\in \mathcal E,\\
&c_i(\mathbf w) ≤ 0, &∀i\in \mathcal I,
\end{aligned}
\tag{NL}
$$
where $\mathcal E\subset ℕ_0$ and $\mathcal I \subset ℕ_0$ index
the nonlinear equality and inequality constraints.
The so called **Lagrangian** $\mathcal L$ of this problem is
$$
\mathcal L(\mathbf w, \mathbf λ, \mathbf μ)
=
L(\mathbf w) + \sum_{i \in \mathcal I} λ_i c_i(\mathbf w) + \sum_{i \in \mathcal E} μ_i c_i(\mathbf w).
$$
Under certain conditions, the Lagrangian can be used to give
necessary conditions for a point $\mathbf w^*$ to be optimal for
(NL).
If $\mathbf w^*$ is locally optimal, then there are multipliers $\mathbf λ^*$
and $\mathbf μ^*$ such that
$$
\begin{aligned}
∇_{\mathbf w} \mathcal L(\mathbf w^*, \mathbf λ^*, \mathbf μ^*) &= \mathbf 0,  \\
c_i(\mathbf w^*) &= 0, & ∀i \in \mathcal E, \\
c_i(\mathbf w^*) &≤ 0, & ∀i \in \mathcal I, \\
λ_i &\geq 0, & ∀i\in \mathcal I,\\
λ_i c_i(\mathbf w^*) &= 0 & ∀i\in \mathcal I.\\
\end{aligned}
$$
These conditions are stated as Theorem 12.1 in “Numerical Optimization” by Nocedal,
and are referred to as the KKT conditions.
The last line, the complementary slackness condition,
tells us, that the multipliers of the inactive
inequality constraints ($c_i(\mathbf w^*) < 0$) are zero.
The first line can be re-written as
$$
\nabla_{\mathbf w}
\mathcal L(\mathbf w^*, \mathbf λ^*, \mathbf μ^*)
=
∇L(\mathbf w)
+
\sum_{i \in \mathcal A(\mathbf w^*)}
λ_i \nabla c_i(\mathbf w^*)
+
\sum_{i\in \mathcal E}
μ_i \nabla c_i(\mathbf w^*)
=
\mathbf 0,
$$
with the active constraint indices $\mathcal A (\mathbf w^*) = \{i \in \mathcal I: c_i(\mathbf w^*) = 0\}$.

## Active-Set Method for Convex QP

An important building block for solving (NL) with
**sequential quadratic programming** (SQP) is the
minimization of the convex quadratic problem
$$
\begin{aligned}
&\min_{\mathbf p \in ℝ^q}
  \frac{1}{2} \mathbf p^T \mathbf H \mathbf p
+
  \mathbf g^T \mathbf p,
&\text{subject to}\\
&\mathbf c_i^T \mathbf p + d_i = 0,  &∀i \in \mathcal E, \\
&\mathbf c_i^T \mathbf p + d_i ≤ 0,  &∀i \in \mathcal I,
\end{aligned}
\tag{QP}
$$
where $\mathbf H$ is a square $q\times q$ matrix and positive semi-definite,
and $\mathbf g$ is a vector in $ℝ^q$.
$\mathcal E\subset ℕ_0$ is the index set for equality constraints,
$\mathcal I\subset ℕ_0$ is the index set for inequality constraints,
and $\mathbf c_i, i\in \mathcal I\cup \mathcal E,$ are vectors in $ℝ^q$.

To solve (NL) using SQP methods, we take an initial guess $\mathbf w^{(0)}$ and solve problems of the form (QP) to
obtain a step $\mathbf p^{(\ell)}$, resulting in updates $\mathbf w^{(\ell+1)} = \mathbf w^{(\ell)} + \mathbf p^{(\ell)}$.
The matrix $\mathbf H$ is the Hessian of $\mathcal L(\mathbf w^{(\ell)})$ and $\mathbf g$ is the loss gradient at $\mathbf w^{(\ell)}$.

But for now, we focus on just solving (QP).
Best forget about the outer iteration indices $(\ell)$ …
Given a concrete instance of (QP), we want to find optimal $\mathbf p^*$ using the Active-Set method.
The Active-Set method is an iterative scheme in which we start with an initial guess
$\mathbf p^{(0)} \in ℝ^q$
and again solve a sequence of sub-problems to determine some step
$\mathbf s^{(k)} = \mathbf p^{(k+1)} - \mathbf p^{(k)}$.
Actually, we have a two-stage procedure.
First, we obtain $\tilde {\mathbf s}^{(k)}$ as the optimizer of
$$
\begin{aligned}
&\min_{\mathbf s \in ℝ^q}
\frac{1}{2} \mathbf s^T \mathbf H \mathbf s
+ \mathbf s^T \mathbf h^{(k)}
&\text{s.t.}\\
&\mathbf c_i^T \mathbf s = 0,  &∀i \in \mathcal E, \\
&\mathbf c_i^T \mathbf s = 0,  &∀i \in \mathcal A^{(k)}.
\end{aligned}
\tag{SP}
$$
Here, we have defined
$$
\mathbf h^{(k)} = \mathbf H \mathbf p^{(k)} + \mathbf g.
$$
The index set $\mathcal A^{(k)} \subset \mathcal I$ is the
iteration-dependent *working set* of active constraints at $\mathbf p^{(k)}$:
$$
\mathcal A^{(k)} = \left\{
   i \in \mathcal I: \mathbf c_i^T \mathbf p^{(k)} + d_i = 0
\right\}.
$$
In a second step, we compute a step-size (Exercise 1c) to obtain update step $\mathbf s^{(k)}$ from $\tilde{\mathbf s}^{(k)}$.

### Exercise 1a) 
In the cell below, complete the code to implement a function, that identifies the active
indices of the linear inequality constraints.
(That is, the inequality value is between 0 and `tol`.)
Suppose, the linear inequality constraints are given in matrix form by a matrix `C_ineq`
and `d_ineq`, such that
$$
\mathbf C_{\text{ineq}} \mathbf p^{(k)} + \mathbf d_{\text{ineq}} \stackrel{!}{≤} \mathbf 0,
$$
i.e., the rows of `C_ineq` hold the constraint vectors, and we are interested in
a subset of the row indices. `active_indices` should return a `Vector{Int}` holding these indices.

<!-- grade_id: cell-1f9babaa1d8fe909, points: 0, grade: false, solution: true -->
```julia
function active_indices(C_ineq, d_ineq, pk, pthrough=nothing; tol=1e-10)
    # ignore the `pthrough` argument. it is used for grading.
    
    @assert size(C_ineq, 1) == length(d_ineq)

    # 1) Initialize empty `Int` vector to hold active indices.
    ### BEGIN SOLUTION
    Ak = Int[]
    ### END SOLUTION
    
    # 2) Evaluate constraint equation C * p + d
    ### BEGIN SOLUTION
    r = C_ineq * pk + d_ineq
    ### END SOLUTION
    
    # 3) Iterate over result vector entries
    #    a) if a value is > `tol`, throw an error, e.g., `error("pk is not feasible")`
    #    b) if *absolute* value of entry is <= `tol`, add index to index array
    ### BEGIN SOLUTION
    for (i, ri) = enumerate(r)
        # if any entry is positive, throw an error. `pk` is not feasible.
        (ri > tol) && error("`pk` is not feasible.")

        # if absolute value of `ri` is <= tol, add `ri` to `Ak`
        if abs(ri) <= tol
            push!(Ak, i)
        end
    end
    ### END SOLUTION

    # 4) return index array
    ### BEGIN SOLUTION
    return Ak
    ### END SOLUTION
end
```

Output:
```
active_indices (generic function with 2 methods)
```

When checking `active_indices`, we can use the `Test` library to test for exceptions:

<!-- grade_id: cell-512477ced7c8ab5c, points: 1, grade: true, solution: false -->
```julia
import Test
let
    # w1 ≤ 0.5*w2
    # w1 ≤ 1
    C_ineq = [
        1 -0.5;
        1 0
    ]
    d_ineq = [0, -1]
    pk = [0.5, 1]
    Ak = active_indices(C_ineq, d_ineq, pk)
    @assert Ak isa AbstractVector{<:Integer}
    @assert only(Ak) == 1

    pk = [2, 1]
    Test.@test_throws Exception active_indices(C_ineq, d_ineq, pk)
    ### BEGIN HIDDEN TESTS
    pk = [1, 2]
    Ak = active_indices(C_ineq, d_ineq, pk)
    @assert Ak == [1, 2]

    pk = [1 - 1e-11, 2]
    Ak = active_indices(C_ineq, d_ineq, pk; tol=1e-10)
    @assert Ak == [1, 2]
    Ak = active_indices(C_ineq, d_ineq, pk; tol=1e-11)
    @assert isempty(Ak)
    ### END HIDDEN TESTS
end
```

### Exercise 1b) 
In the cell below, complete the code to implement a function, that solves the KKT
equations for the subproblem (SP),
$$
\begin{bmatrix}
\mathbf H & \mathbf C_{\mathcal A^{(k)}}^T & \mathbf C_{\text{eq}}^T \\
\mathbf C_{\mathcal A^{(k)}}& \mathbf 0 & \mathbf 0 \\
\mathbf C_{\text{eq}} & \mathbf 0 & \mathbf 0
\end{bmatrix}
\begin{bmatrix}
\tilde{\mathbf s}^{(k)} \\
\mathbf λ^{(k)} \\
\mathbf μ^{(k)} \\
\end{bmatrix}
\stackrel{!}{=}
\begin{bmatrix}
-\mathbf h^{(k)} \\
\mathbf 0\\
\mathbf 0
\end{bmatrix}.
$$
These equations result from writing down the KKT conditions for (SP).

<!-- grade_id: cell-6230ae4ebc873df2, points: 0, grade: false, solution: true -->
```julia
function solve_step_problem(H, C_Ak, C_eq, hk, pthrough=nothing)
    # ignore the `pthrough` argument. it is used for grading.

    num_vars = size(H, 1)  # length of sk_tilde
    @assert size(H, 2) == num_vars
    @assert size(C_Ak, 2) == num_vars
    @assert size(C_eq, 2) == num_vars
    @assert length(hk) == num_vars

    num_ineq = size(C_Ak, 1) # length of lambdak
    num_eq = size(C_eq, 1)   # length of muk

    # change `KKT` to hold the correct matrix
    KKT = rand(num_vars + num_ineq + num_eq, num_vars + num_ineq + num_eq)
    ### BEGIN SOLUTION
    KKT = [
        H C_Ak' C_eq';
        C_Ak zeros(num_ineq, num_ineq) zeros(num_ineq, num_eq);
        C_eq zeros(num_eq, num_ineq) zeros(num_eq, num_eq)
    ]
    ### END SOLUTION
    
    RHS = zeros(num_vars + num_ineq + num_eq)
    # change `RHS` to hold the correct right-hand-side matrix
    ### BEGIN SOLUTION
    RHS[1:num_vars] .= -hk
    ### END SOLUTION

    # solve the linear equation system and set `sk`, `lambdak` and `muk`
    # to hold the multipliers
    sk_tilde = lambdak = muk = nothing
    ### BEGIN SOLUTION
    multipliers = KKT \ RHS

    sk_tilde = @view(multipliers[1:num_vars])
    lambdak = @view(multipliers[num_vars+1:num_vars+num_ineq])
    muk = @view(multipliers[num_vars+num_ineq+1:end])
    ### END SOLUTION
    return sk_tilde, lambdak, muk
end
```

Output:
```
solve_step_problem (generic function with 2 methods)
```

Below are the tests for `solve_step_problem`. We use `LinearAlgebra` to initialize `H` as the identity matrix.

```julia
import LinearAlgebra as LA
```

<!-- grade_id: cell-1e63ab26bb8ea888, points: 1, grade: true, solution: false -->
```julia
let
    H = 1.0 .* LA.I(2)
    C_Ak = [1 -0.5]
    C_eq = [1 0]
    pk = [1, 2]
    hk = H * pk
    sk_tilde, λk, μk = solve_step_problem(H, C_Ak, C_eq, hk)

    @assert sk_tilde isa AbstractVector
    @assert λk isa AbstractVector
    @assert μk isa AbstractVector
    @assert length(sk_tilde) == 2
    @assert length(λk) == 1
    @assert length(μk) == 1
    # `sk_tilde` should be close to zero:
    @assert sum(sk_tilde .^ 2) <= 1e-10
    ### BEGIN HIDDEN TESTS
    @assert only(λk) ≈ 4
    @assert only(μk) ≈ -5

    C_Ak = zeros(0, 2)
    pk = [1, 4]
    hk = H * pk
    sk_tilde, λk, μk = solve_step_problem(H, C_Ak, C_eq, hk)
    @assert sk_tilde ≈ [0, -4]
    @assert isempty(λk)
    @assert μk ≈ [-1]
    ### END HIDDEN TESTS
    nothing
end
```

### Exercise 1c)
Once $\tilde{\mathbf s}^{(k)}$ is available from `solve_step_problem`,
we need to determine a step size $α_k \in [0, 1]$.
$α_k$ should be as large as possible, and chosen so
that $\mathbf p^{(k+1)} = \mathbf p^{(k)} + α_k \tilde{\mathbf s}^{(k)}$ is
feasible for **all** the linear constraints of (QP).
The vector $\mathbf p^{(k+1)}$ is feasible for any $α_k$ and
all the constraints indexed by $\mathcal E$
and $\mathcal A^{(k)}$ by virtue of the KKT conditions of (SP).
Suppose $i\in \mathcal I\smallsetminus A^{(k)}$.
If $\mathbf c_i^T \tilde{\mathbf s}^{(k)} \leq 0$, then for all $α_k \geq 0$
we have
$$
\mathbf c_i^T(\mathbf p^{(k)} + α_k \tilde{\mathbf s}^{(k)}) + d_i
\leq
\mathbf c_i^T \mathbf p^{(k)} + d_i
\leq 0.
$$
Thus, we only have to consider the case
$i \in \mathcal I\smallsetminus \mathcal A^{(k)}$
and $\mathbf c_i^T \tilde{\mathbf s}^{(k)} > 0$.
In that case, we solve
$$
α_k
\leq
\frac{-d_i - \mathbf c_i^T \mathbf p^{(k)}}{\mathbf c_i^T \tilde{\mathbf s}^{(k)}}.
$$
Implement `stepsize_and_blocking_index`
to return the stepsize $α_k$ and a blocking index.
Use this formula:
$$
α_k = \min
\left\{
1,
\min_{
    \substack{
        i \in \mathcal I\smallsetminus \mathcal A^{(k)},\\
        \mathbf c_i^T \tilde{\mathbf s}^{(k)} > 0
    }
}
\frac{-d_i - \mathbf c_i^T \mathbf p^{(k)}}{\mathbf c_i^T \tilde{\mathbf s}^{(k)}}
\right\}.
$$
A "blocking index" is an index $i\in\mathcal I\smallsetminus \mathcal A^{(k)}$ such that the expression
of the inner $\min$ operator is strictly smaller than $1$.
If there is no blocking index, return the stepsize and $-1$.

<!-- grade_id: cell-b0981f4b00c0721c, points: 0, grade: false, solution: true -->
```julia
function stepsize_and_blocking_index(C_ineq, d_ineq, pk, sk_tilde, Ak, pthrough=nothing)
    # ignore the `pthrough` argument. it is used for grading.

    alphak = 1.0
    j = -1

    # determine correct stepsize `alphak` according to the formula above;
    # `sk_tilde` is the unscaled step vector, `pk` are the variables,
    # the rows of `C_ineq` hold the vectors $c_i$,
    # and the vector `d_ineq` holds the linear constraint offsets:
    ### BEGIN SOLUTION
    # indices that are **not** active:
    Mk = setdiff(axes(C_ineq, 1), Ak)
    # extract partial matrices
    C = @view(C_ineq[Mk, :])
    d = @view(d_ineq[Mk])

    # compute matrix-vector products beforehand (optional)
    Cs = C * sk_tilde
    r = C * pk .+ d
    # loop over all non-active constraints and seek minimum stepsize value
    for i in axes(Cs, 1)
        Cs[i] <= 0 && continue
        alpha = -r[i] / Cs[i]
        if alpha <= alphak
            alphak = alpha
            j = Mk[i]
        end
    end
    ### END SOLUTION
    return alphak, j
end
```

Output:
```
stepsize_and_blocking_index (generic function with 2 methods)
```

Let's test our function:

<!-- grade_id: cell-265c0df05549a1b2, points: 1.5, grade: true, solution: false -->
```julia
let
    C_ineq = [1 -0.5]
    d_ineq = [0,]

    pk = [1.0, 2.0]
    sk_tilde = [0.0; -2.0]
    Ak = Int[]
    # Check if the returned values make sense:
    αk, j = stepsize_and_blocking_index(C_ineq, d_ineq, pk, sk_tilde, Ak)
    @assert αk isa Real
    @assert j isa Integer
    @assert -1e-10 <= αk <= 1 + 1e-10
    @assert j >= -1
    ### BEGIN HIDDEN TESTS
    @assert abs(αk) <= 1e-10
    @assert j == 1
    ### END HIDDEN TESTS

    # Now, we should be able to take a full step:
    Ak = [1,]
    sk = zeros(2)

    αk, j = stepsize_and_blocking_index(C_ineq, d_ineq, pk, sk_tilde, Ak)
    @assert αk ≈ 1
    @assert j == -1
    ### BEGIN HIDDEN TESTS
    C_ineq = [1 -0.5; 10 100]
    d_ineq = [0, 1000]
    Ak = Int[]
    pk = [1.0, 4.0]
    sk_tilde = [0, -4]
    αk, j = stepsize_and_blocking_index(C_ineq, d_ineq, pk, sk_tilde, Ak)
    @assert αk ≈ 0.5
    @assert j == 1
    ### END HIDDEN TESTS
end
```

### Exercise 1d) 
We now have all building blocks for the Active-Set Method.
The algorithm stops, if a step $\mathbf s^{(k)}$ is zero
and all multipliers $λ_i^{(k)}$ are non-negative.
Otherwise, the working set is augmented.
Likewise, the working set is augmented in case of a non-zero
step with $α_k < 1$.
Complete the marked sections in the code below:

<!-- grade_id: cell-f474efae4b45a7ff, points: 0, grade: false, solution: true -->
```julia
function solve_convex_qp(
    H, g, C_eq, d_eq, C_ineq, d_ineq, p0, pthrough=nothing;
    tol=1e-10, max_iter=100
)
    # ignore the `pthrough` argument. it is used for grading.

    
    pk = copy(p0)
    # 1) check feasibilty of p0 with respect to equality constraints:
    #    compute vector `r_eq` holding the equality constraint values and
    #    check if all entries have absolute values <= tol; elsewise `error`.
    ### BEGIN SOLUTION
    r_eq = C_eq * pk .+ d_eq
    if any(abs.(r_eq) .> tol)
        error("`p0` is not feasible for equality constraints.")
    end
    ### END SOLUTION
    
    # 2) determine active indices for `pk` with `active_indices` and store in vector `Ak`
    # IMPORTANT: Pass `pthrough` to `active_indices`!
    Ak = Int[]
    ### BEGIN SOLUTION
    Ak = active_indices(C_ineq, d_ineq, pk, pthrough; tol)
    ### END SOLUTION

    k = 1
    while k <= max_iter
        # Compute the current objective value `Lk` of (QP) at `pk`
        Lk = Inf
        ### BEGIN SOLUTION
        Lk = 0.5 * pk'H * pk + g'pk
        ### END SOLUTION

        @info "Iteration $k, L(pk)=$(Lk)."

        # extract active inequality constraint matrix
        C_Ak = @view(C_ineq[Ak, :])

        # solve SP to get step:
        # first, compute `hk` and then call `solve_step_problem`;
        # pass the correct constraint matrices and
        # store results in `sk_tilde, λk, μk`.
        # IMPORTANT: Pass `pthrough` to `solve_step_problem`!
        hk = sk_tilde = λk = μk = nothing
        ### BEGIN SOLUTION
        hk = H * pk .+ g
        sk_tilde, λk, μk = solve_step_problem(H, C_Ak, C_eq, hk, pthrough)
        ### END SOLUTION

        if sqrt(sum(sk_tilde .^ 2)) <= tol
            @info "Zero step."
            # if `sk_tilde` is (nearly) zero, then see if all multipliers are non-negative
            problem_solved = isempty(λk)
            if !problem_solved
                λ_min, λ_min_index = findmin(λk)
                problem_solved = λ_min >= -tol
            end
            if problem_solved
                @info "Problem solved."
                break
            end
            # if there are negative multipliers,
            # remove the index corresponding to smallest index for next iteration
            deleteat!(Ak, λ_min_index)
        else
            # `sk_tilde` is different from zero, compute stepsize `αk` and see if there
            # is a blocking index `j`.
            # If there is a blocking index, add it to `Ak` for the next iteration.
            # IMPORTANT: Pass `pthrough` to `stepsize_and_blocking_index`!
            ### BEGIN SOLUTION
            αk, j = stepsize_and_blocking_index(C_ineq, d_ineq, pk, sk_tilde, Ak)
            @info "Steplength is $(αk)."
            if j > 0
                @info "Adding $j to working set."
                push!(Ak, j)
            end
            ### END SOLUTION
            # Also, update `pk` for the next iteration:
            ### BEGIN SOLUTION
            pk .+= αk .* sk_tilde
            ### END SOLUTION
        end
        k += 1
    end
    return pk, k
end
```

Output:
```
solve_convex_qp (generic function with 2 methods)
```

Here is how we can test your implementation:

```julia
pthrough=nothing
H = 1.0 .* LA.I(2)
g = zeros(2)

C_ineq = [1 -0.5]
d_ineq = [0,]

C_eq = [1 0]
d_eq = [-1,]

p0 = [1, 4]
popt, k = solve_convex_qp(H, g, C_eq, d_eq, C_ineq, d_ineq, p0, pthrough; max_iter=10)
@assert popt isa Vector
@assert length(popt) == 2
@assert k == 2
# check that solution satisfies constraints
@assert abs(popt[1] - 1) <= 1e-9
@assert abs(popt[1] - 0.5 * popt[2]) <= 1e-9
# check, if optimal value is reduced
L(p) = 0.5 * p'H * p + g'p
L0 = L(p0)
Lopt = L(popt)
@assert Lopt < L0
```

Output:
```
[36m[1m[ [22m[39m[36m[1mInfo: [22m[39mIteration 1, L(pk)=8.5.
[36m[1m[ [22m[39m[36m[1mInfo: [22m[39mSteplength is 0.5.
[36m[1m[ [22m[39m[36m[1mInfo: [22m[39mAdding 1 to working set.
[36m[1m[ [22m[39m[36m[1mInfo: [22m[39mIteration 2, L(pk)=2.5.
[36m[1m[ [22m[39m[36m[1mInfo: [22m[39mZero step.
[36m[1m[ [22m[39m[36m[1mInfo: [22m[39mProblem solved.
```

```julia
# NBGrader does not like logging, so we disable it some times:
import Logging
```

<!-- grade_id: cell-c6f988b7a30df7d2, points: 1.5, grade: true, solution: false -->
```julia
let pthrough = Val(true);
    Logging.with_logger(Logging.NullLogger()) do
        H = 1.0 .* LA.I(2)
        g = zeros(2)

        C_ineq = [1 -0.5]
        d_ineq = [0,]

        C_eq = [1 0]
        d_eq = [-1,]

        p0 = [1, 4]
        popt, k = solve_convex_qp(H, g, C_eq, d_eq, C_ineq, d_ineq, p0, pthrough; max_iter=10)
        @assert popt isa Vector
        @assert length(popt) == 2
        @assert k == 2
        # check that solution satisfies constraints
        @assert abs(popt[1] - 1) <= 1e-9
        @assert abs(popt[1] - 0.5 * popt[2]) <= 1e-9
        # check, if optimal value is reduced
        L(p) = 0.5 * p'H * p + g'p
        L0 = L(p0)
        Lopt = L(popt)
        @assert Lopt < L0

        ### BEGIN HIDDEN TESTS
        p0 = [1, 2]
        popt, k = solve_convex_qp(H, g, C_eq, d_eq, C_ineq, d_ineq, p0, pthrough; max_iter=10)
        @assert popt ≈ [1, 2]
        @assert k == 1
        ### END HIDDEN TESTS
    end
end
```

In the cell below, we can see if you successfully disallow infeasible points:

<!-- grade_id: cell-29df28bd7c0c2b03, points: 0.5, grade: true, solution: false -->
```julia
let
    # same setup as before:
    H = 1.0 .* LA.I(2)
    g = zeros(2)
    C_ineq = [1 -0.5]
    d_ineq = [0,]
    C_eq = [1 0]
    d_eq = [-1,]

    p0 = [2, 1]

    Test.@test_throws Exception solve_convex_qp(H, g, C_eq, d_eq, C_ineq, d_ineq, p0; max_iter=10)
    ### BEGIN HIDDEN TESTS
    p0 = [0.5, 2]
    Test.@test_throws Exception solve_convex_qp(H, g, C_eq, d_eq, C_ineq, d_ineq, p0; max_iter=10)
    ### END HIDDEN TESTS
end
```

Output:
```
[32m[1mTest Passed[22m[39m
      Thrown: ErrorException
```

## JuMP for Linear and Quadratic Programming

The [`JuMP` package](https://jump.dev/JuMP.jl/stable/) provides a unified syntax to
define problems with linear or quadratic constraints.
It also supports automatic re-formulation of problems to suit the needs of
the various solvers that are supported.
[The list of solvers](https://jump.dev/JuMP.jl/stable/installation/#Supported-solvers) is
pretty extensive, and the ability to switch between them so easily is one of the
main reasons why `JuMP` is very popular with practitioners in the fields of optimization
and optimal control.

If we were to write an actual SQP algorithm, and did not want to bother with
the Active-Set-Method subroutine ourselves, we could use `JuMP` to solve the sub-problems.
In exercise 2a) you are going to use `JuMP` and the solver `COSMO` to solve the following
problem:
$$
\begin{alignedat}{3}
&\min_{\mathbf p\in ℝ^2}
   (p_1 - 1)^2 + (p_2 - 2.5)^2
&& &&\text{subject to}\\
&-p_1 + 2 p_2 - 2 &&≤ 0,\\
&+p_1 + 2 p_2 - 6 &&≤ 0,\\
&+p_1 - 2 p_2 - 2 &&≤ 0,\\
&+p_1 && \geq 0,\\
&+p_2 && \geq 0.\\
\end{alignedat}
\tag{P}
$$

![problem plot](fig_cosmo.png)

### Exercise 2a) 
Below is a function to solve the problem (P) using `JuMP`.
However, there are 5 mistakes in the code.
Fix them. If unsure, consult the [Tutorial.](https://jump.dev/JuMP.jl/stable/tutorials/getting_started/getting_started_with_JuMP/#Getting-started-with-JuMP)

```julia
import JuMP
import JuMP: @variable, @objective, @constraint, optimize!, Model, value
import COSMO
```

<!-- grade_id: cell-d8fe6ab0b20f826c, points: 0, grade: false, solution: true -->
```julia
p_cosmo = model = nothing
let
    # we are using a local scope here to be extra careful, and to avoid poluting the
    # global scope
    global p_cosmo, model
    model = Model(COSMO.Optimizer)
    @variable(model, p[1:3] .>= -1)

    @objective(model, Min, (p[1] - 1)^2 + (p[2] - 2.5)^2 + p[3])

    @constraint(model, -p[1] + 2 * p[2] + 2 <= 0)
    @constraint(model, p[1] + 2 * p[2] - 6 <= 0)
    @constraint(model, +p[1] - 2 * p[2] - 2 <= 2)
    ### BEGIN SOLUTION
    model = Model(COSMO.Optimizer)
    @variable(model, p[1:2] .>= 0)
    @objective(model, Min, (p[1] - 1)^2 + (p[2] - 2.5)^2)
    @constraint(model, -p[1] + 2 * p[2] - 2 <= 0)
    @constraint(model, p[1] + 2 * p[2] - 6 <= 0)
    @constraint(model, +p[1] - 2 * p[2] - 2 <= 0)
    ### END SOLUTION

    optimize!(model)

    @show p_cosmo = value.(p)
    nothing
end
```

Output:
```
------------------------------------------------------------------
          COSMO v0.8.9 - A Quadratic Objective Conic Solver
                         Michael Garstka
                University of Oxford, 2017 - 2022
------------------------------------------------------------------

Problem:  x ∈ R^{2},
          constraints: A ∈ R^{5x2} (8 nnz),
          matrix size to factor: 7x7,
          Floating-point precision: Float64
Sets:     Box of dim: 3
          Nonnegatives of dim: 2
Settings: ϵ_abs = 1.0e-05, ϵ_rel = 1.0e-05,
          ϵ_prim_inf = 1.0e-04, ϵ_dual_inf = 1.0e-04,
          ρ = 0.1, σ = 1e-06, α = 1.6,
          max_iter = 5000,
          scaling iter = 10 (on),
          check termination every 25 iter,
          check infeasibility every 40 iter,
          KKT system solver: COSMO.QdldlKKTSolver
Acc:      Anderson Type2{QRDecomp},
          Memory size = 7, RestartedMemory,	
          Safeguarded: true, tol: 2.0
Setup Time: 1694.44ms

Iter:	Objective:	Primal Res:	Dua
```

Check success:

<!-- grade_id: cell-d65ee4b485b82aa6, points: 0.5, grade: true, solution: false -->
```julia
@assert JuMP.termination_status(model) == JuMP.OPTIMAL
```

Check dimensions and constraint types:

<!-- grade_id: cell-ecb8520065f6c735, points: 0.5, grade: true, solution: false -->
```julia
@assert length(p_cosmo) == 2
@assert JuMP.num_constraints(model; count_variable_in_set_constraints=true) == 5
```

Check optimal value:

<!-- grade_id: cell-a7231d69f0a90539, points: 1, grade: true, solution: false -->
```julia
let p = p_cosmo
    @assert all(p .>= 0)
    @assert -p[1] + 2 * p[2] - 2 <= 1e-10
    @assert p[1] + 2 * p[2] - 6 <= 1e-10
    @assert +p[1] - 2 * p[2] - 2 <= 1e-10
    ### BEGIN HIDDEN TESTS
    @assert p ≈ [1.4, 1.7]
    ### END HIDDEN TESTS
end
```

![problem plot](fig_cosmo2.png)

### Exercise 2b) 
As indicated, it is possible to use `JuMP` in many different ways.
We could also provide the affine-linear inequality constraints with matrices.
Set the 3 x 2 matrix `A_cosmo` and the 3-vector `b_cosmo`, so that `A_cosmo * x .<= b_cosmo`
describes the constraints of (P) (other than the non-negativity constraints.)

<!-- grade_id: cell-3d8f2b57ef65c61f, points: 0, grade: false, solution: true -->
```julia
A_cosmo = zeros(3, 2)
b_cosmo = zeros(3)
### BEGIN SOLUTION
A_cosmo = [
    -1 2;
    1 2;
    1 -2
]
b_cosmo = [2, 6, 2]
nothing
### END SOLUTION
```

Moreover, we can also rewrite the objective function as
`0.5 * p' * H_cosmo * p + p'g_cosmo + c_cosmo`,
with a 2 x 2 matrix `H_cosmo` and vector `g_cosmo` and constant `c_cosmo`:

<!-- grade_id: cell-c55f2768af3d5d6b, points: 0, grade: false, solution: true -->
```julia
H_cosmo = zeros(2, 2)
g_cosmo = zeros(2)
c_cosmo = 0
### BEGIN SOLUTION
# (p1 - 1)^2 = p1^2 - 2*p1 + 1
# (p2 - 2.5)^2 = p2^2 - 5*p2 + 2.5^2
# objf == p1^2 + p2^2 - 2*p1 - 5*p2 + 1 + 2.5^2
H_cosmo = [2 0; 0 2]
g_cosmo = [-2, -5]
c_cosmo = 2.5^2 + 1
nothing
### END SOLUTION
```

Check against last optimal value:

<!-- grade_id: cell-280b229a7296a2f9, points: 1, grade: true, solution: false -->
```julia
@assert JuMP.objective_value(model) ≈ 0.5 * p_cosmo'H_cosmo * p_cosmo + p_cosmo'g_cosmo + c_cosmo
```

Here is how to use matrices with `JuMP`:

```julia
let
    global H_cosmo, g_cosmo, c_cosmo
    global A_cosmo, b_cosmo, p_cosmo
    model = Model(COSMO.Optimizer)
    @variable(model, p[1:2] .>= 0)
    @objective(model, Min, 0.5 * p'H_cosmo * p + p'g_cosmo + c_cosmo)
    @constraint(model, A_cosmo * p .<= b_cosmo)
    optimize!(model)

    p_opt = value.(p)
    @assert p_opt ≈ p_cosmo
end
```

Output:
```
------------------------------------------------------------------
          COSMO v0.8.9 - A Quadratic Objective Conic Solver
                         Michael Garstka
                University of Oxford, 2017 - 2022
------------------------------------------------------------------

Problem:  x ∈ R^{2},
          constraints: A ∈ R^{5x2} (8 nnz),
          matrix size to factor: 7x7,
          Floating-point precision: Float64
Sets:     Box of dim: 3
          Nonnegatives of dim: 2
Settings: ϵ_abs = 1.0e-05, ϵ_rel = 1.0e-05,
          ϵ_prim_inf = 1.0e-04, ϵ_dual_inf = 1.0e-04,
          ρ = 0.1, σ = 1e-06, α = 1.6,
          max_iter = 5000,
          scaling iter = 10 (on),
          check termination every 25 iter,
          check infeasibility every 40 iter,
          KKT system solver: COSMO.QdldlKKTSolver
Acc:      Anderson Type2{QRDecomp},
          Memory size = 7, RestartedMemory,	
          Safeguarded: true, tol: 2.0
Setup Time: 0.2ms

Iter:	Objective:	Primal Res:	Dual Re
```

### Exercise 2c) 
Incidently, we have restated our problem in a form nearly suitable for our Active-Set solver
from exercise 1.
Augment `A_cosmo` and `b_cosmo` by including the non-negativity constraints and
name the resulting matrix `C_ineq_as`, and the vector `b_ineq_as`:

<!-- grade_id: cell-8f445681bfb6c154, points: 0, grade: false, solution: true -->
```julia
C_ineq_as = zeros(5, 2)
b_ineq_as = zeros(5)
### BEGIN SOLUTION
C_ineq_as = vcat([-1 0; 0 -1], A_cosmo)
b_ineq_as = vcat(zeros(2), b_cosmo)
nothing
### END SOLUTION
```

Finally, our solver should return the same optimum, hopefully:

```julia
let H = H_cosmo, g = g_cosmo, C_ineq = C_ineq_as, d_ineq = -b_ineq_as
    global p_cosmo
    C_eq = zeros(0, 2)
    d_eq = []
    p0 = ones(2)
    p_qp = p0
    p_qp, _ = solve_convex_qp(H, g, C_eq, d_eq, C_ineq, d_ineq, p0)

    @assert p_cosmo ≈ p_qp
end
```

Output:
```
[36m[1m[ [22m[39m[36m[1mInfo: [22m[39mIteration 1, L(pk)=-5.0.
[36m[1m[ [22m[39m[36m[1mInfo: [22m[39mSteplength is 0.3333333333333333.
[36m[1m[ [22m[39m[36m[1mInfo: [22m[39mAdding 3 to working set.
[36m[1m[ [22m[39m[36m[1mInfo: [22m[39mIteration 2, L(pk)=-6.25.
[36m[1m[ [22m[39m[36m[1mInfo: [22m[39mSteplength is 1.0.
[36m[1m[ [22m[39m[36m[1mInfo: [22m[39mIteration 3, L(pk)=-6.450000000000001.
[36m[1m[ [22m[39m[36m[1mInfo: [22m[39mZero step.
[36m[1m[ [22m[39m[36m[1mInfo: [22m[39mProblem solved.
```